# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In this GitHub repository, there is a file called `allergies.json` that contains a list of patient allergies.  You will need to download this [file from here](https://raw.githubusercontent.com/paulboal/hds5210-2023/main/week06/allergies.json) and then upload it into Google Colab to run these examples. It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE=json.load(open('allergies.json'))
type(ALLERGIES_FILE)
```

Output:
```
dict
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_count(json_file):
  """
   (dict)->int
   Here i have created a varible as events and I have run the for loop so that i
   loop over the enteries an I have put the if condition to access the certainity
   from the entry so if the certainity is confirmed it adds up the events so that
   we get to see how many allergies are recorded.

   >>> (allergy_count(ALLERGIES_FILE))
   int

   >>> allergy_count(ALLERGIES_FILE)
   4
  """
  try:
    with open(json_file,'r') as file:
      data=json.load(file)
      if "entry" in data and isinstance(data["entry"],list):
        return len(data["entry"])
      else:
        return 0
  except FileNotFoundError:
    return 0
ALLERGIES_FILE='allergies.json'
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

```python
import doctest
doctest.run_docstring_examples(allergy_count,globals(),verbose=True)
```

Output:
```
PYDEV DEBUGGER WARNING:
sys.settrace() should not be used when the debugger is being used.
This may cause the debugger to stop working correctly.
If this is needed, please check: 
http://pydev.blogspot.com/2007/06/why-cant-pydev-debugger-work-with.html
to see how to restore the debug tracing back correctly.
Call Location:
  File "/usr/lib/python3.10/doctest.py", line 1501, in run
    sys.settrace(save_trace)

Finding tests in NoName
Trying:
    (allergy_count(ALLERGIES_FILE))
Expecting:
    int
**********************************************************************
File "__main__", line 10, in NoName
Failed example:
    (allergy_count(ALLERGIES_FILE))
Expected:
    int
Got:
    4
Trying:
    allergy_count(ALLERGIES_FILE)
Expecting:
    4
ok
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
# Put your solution here
import json

def patient_count(json_file):
  """
  (dict)->int
   Here I have created a variable as events and I have run the for loop so that i
   loop over the enteries an I have put the if conditions to access the certainity
   from the entry so if the certainity is confirmed it adds up the event so that
   we get to see how many allergies are recorded.

   >>> type(patient_count(ALLERGIES_FILE))
   int

   >>> patient_count(ALLERGIES_FILE)
    2
    """
  try:
      with open(json_file, 'r') as file:
          data = json.load(file)
          unique_patients = set()

          if "entry" in data and isinstance(data["entry"], list):
              for entry in data["entry"]:
                   if "resource" in entry and "patient" in entry["resource"] and "display" in entry["resource"]["patient"]:
                      patient_name = entry["resource"]["patient"]["display"]
                      unique_patients.add(patient_name)

          return len(unique_patients)
  except FileNotFoundError:
        return 0  # Return 0 if the file is not found or there is an error while reading

# Define the path to your JSON file (e.g., ALLERGIES_FILE)
ALLERGIES_FILE = 'allergies.json'
```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
# Put your solution here
import json

def allergy_per_patient(json_file):
  """
     (dict)->int
     Here I have created a variable as events and I have run the for loop so that  i
     loop over the enteries an I have put the if condition to access the certainity
     from the entry so if the certainity is confirmed it adds up the events so that
     we get to see how many allergies are recorded.

     >>> type (allergy_per_patient("ALLERGIES_FILE.json"))
     dict

     >>> allergy_per_patient("ALLERGIES_FILE.json")
     {'Paul Boal':1,'Jason Argonaut':3}
    """
  try:
      with open(json_file, 'r') as file:
          data = json.load(file)
          allergy_counts = {}  # Dictionary to store allergy counts per patient

          if "entry" in data and isinstance(data["entry"], list):
            for entry in data["entry"]:
                  if "resource" in entry and "patient" in entry["resource"]:
                      patient_name = entry["resource"]["patient"]["display"]

                        # Check if the patient_name is already in the dictionary
                      if patient_name in allergy_counts:
                          allergy_counts[patient_name] += 1
                      else:
                          allergy_counts[patient_name] = 1
            print(allergy_counts)
          return allergy_counts
  except FileNotFoundError:
        return {}

# Define the path to your JSON file (e.g., ALLERGIES_FILE)
ALLERGIES_FILE = 'allergies.json'
```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
import json

def allergy_list(json_file):
    """
  (dict)->list
  Here I have created a variable as events and I have run the for loop so that i
  loop over the enteries an I have put the if conditions to access the certainity
  from the entry so if the certainity is confirmed it adds up the event so that
  we get to see how many allergies are recorded.

  >>> type(allergy_list("ALLERGIES_FILE.json"))

   >>> allergy_list(ALLERGIES_FILE)
   [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
   ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
   ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
   ['Paul Boal', 'PENICILLIN G', 'Bruising']]
   """
    try:
        with open(json_file, 'r') as file:
            data = json.load(file)
            output_list = []

            if "entry" in data and isinstance(data["entry"], list):
                for entry in data["entry"]:
                    if "resource" in entry and "patient" in entry["resource"]:
                        patient_name = entry["resource"]["patient"]["display"]

                        # Extract allergy and reaction information
                        if "substance" in entry["resource"] and "text" in entry["resource"]["substance"]:
                            allergy = entry["resource"]["substance"]["text"]
                        else:
                            allergy = None

                        if "reaction" in entry["resource"] and isinstance(entry["resource"]["reaction"], list):
                            first_reaction = entry["resource"]["reaction"][0]["manifestation"][0]["text"] if len(entry["resource"]["reaction"]) > 0 and "manifestation" in entry["resource"]["reaction"][0] else None
                        else:
                            first_reaction = None

                        # Add the extracted information to the output list
                        output_list.append([patient_name, allergy, first_reaction])

            return output_list
    except FileNotFoundError:
        return []

# Define the path to your JSON file (e.g., ALLERGIES_FILE)
ALLERGIES_FILE = 'allergies.json'
```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_reaction(json_file, patient, substance):
    """
  (dict)->str or None
  Here I have created a variable as events and I have run the for loop so that i
  loop over the enteries an I have put the if condition to access the certainity
  from the entry so if the certainity is confirmed it adds up the events so that
  we got to see how many allergies are recorded.

  >>> type(allergy_reaction("ALLERGIES_FILE.json"))

  >>> allergy_reaction(ALLERGIES_FILE,'Jason Argonaut','PENICILLIN G')
  Hives
  >>> allergy_reaction(ALLERGIES_FILE,'Jason Argonaut','SHELLFISH-DERIVED PRODUCTS')
  Itching

  >>> allergy_reaction(ALLERGIES_FILE,'Jason Argonaut','STRAWBERRY')
  Anaphylaxis

  >>> allergy_reaction(ALLERGIES_FILE,'Jason Argonaut','PENICILLIN')
  None

  >>> allergy_reaction(ALLERGIES_FILE,'Paul Boal','PENICILLIN G')
  Bruising
  """
    allergy_data = allergy_list(json_file)

    reaction = None

    for entry in allergy_data:
        entry_patient, entry_substance, entry_reaction =entry
        if entry_patient == patient and entry_substance == substance:
            reaction = entry_reaction
    return reaction

# Define the path to your JSON file (e.g., ALLERGIES_FILE)
ALLERGIES_FILE = 'allergies.json'
```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

